```python
# Initialize Otter
import otter
grader = otter.Notebook("hw06.ipynb")
```

# Homework 6: Probability, Simulation, Estimation, and Assessing Models

Please complete this notebook by filling in the cells provided. Before you begin, execute the previous cell to load the provided tests.

**Helpful Resource:**
- [Python Reference](http://data8.org/su24/reference/): Cheat sheet of helpful array & table methods used in Data 8!
- [Sampling Methods Guide](https://drive.google.com/file/d/1k8uNN0Bg7J0BTgK-0Z5glDNtd_5FTJQl/view?usp=sharing): Guide for the randomization methods.

**Recommended Readings**: 
* [Randomness](https://www.inferentialthinking.com/chapters/09/Randomness.html)
* [Sampling and Empirical Distributions](https://www.inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html)
* [Testing Hypotheses](https://www.inferentialthinking.com/chapters/11/Testing_Hypotheses.html)

## 1. Roulette

```python
# Run this cell to set up the notebook, but please don't change it.

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

A Nevada roulette wheel has 38 pockets and a small ball that rests on the wheel. When the wheel is spun, the ball comes to rest in one of the 38 pockets. That pocket is declared the winner. 

The pockets are labeled 0, 00, 1, 2, 3, 4, ... , 36. Pockets 0 and 00 are green, and the other pockets are alternately red and black. The table `wheel` is a representation of a Nevada roulette wheel. **Note that *both* columns consist of strings.** Below is an example of a roulette wheel!

<img src="roulette_wheel.jpeg" width="330px">

Run the cell below to load the `wheel` table.

```python
wheel = Table.read_table('roulette_wheel.csv', dtype=str)
wheel
```

Output:
```
Pocket | Color
00     | green
0      | green
1      | red
2      | black
3      | red
4      | black
5      | red
6      | black
7      | red
8      | black
... (28 rows omitted)
```

### Betting on Red ###
If you bet on *red*, you are betting that the winning pocket will be red. This bet *pays 1 to 1*. That means if you place a one-dollar bet on red, then:

- If the winning pocket is red, you gain 1 dollar. That is, you get your original dollar back, plus one more dollar.
- If the winning pocket is not red, you lose your dollar. In other words, you gain -1 dollars.

Let's see if you can make money by betting on red at roulette.

**Question 1.** Define a function `dollar_bet_on_red` that takes the name of a color and returns your gain in dollars if that color had won and you had placed a one-dollar bet on red. Remember that the gain can be negative. Make sure your function returns an integer.

*Note:* You can assume that the only colors that will be passed as arguments are red, black, and green. Your function doesn't have to check that.


```python
def dollar_bet_on_red(color):
    roll=np.random.choice(wheel.column("Color"))
    outcome=0
    if roll == "red" == color:
        outcome=outcome+1
    if roll == "black" == color:
        outcome=outcome-1
    if roll == "green" == color:
        outcome=outcome-1
    return outcome
```

```python
grader.check("q1_1")
```

Output:
```
q1_1 results: All test cases passed!
```

Run the cell below to make sure your function is working.

```python
print(dollar_bet_on_red('green'))
print(dollar_bet_on_red('black'))
print(dollar_bet_on_red('red'))
```

Output:
```
0
0
1
```

**Question 2.** Add a column labeled `Winnings: Red` to the table `wheel`. For each pocket, the column should contain your gain in dollars if that pocket won and you had bet one dollar on red. Your code should use the function `dollar_bet_on_red`.

> *Hint:* You should not need a `for` loop for this question, instead try using a table method!

```python
red_winnings = ...
wheel = ...
wheel
```

```python
grader.check("q1_2")
```

### Simulating 10 Bets on Red
Roulette wheels are set up so that each time they are spun, the winning pocket is equally likely to be any of the 38 pockets regardless of the results of all other spins. Let's see what would happen if we decided to bet one dollar on red each round.

**Question 3.** Create a table `ten_bets` by sampling the table `wheel` to simulate 10 spins of the roulette wheel. Your table should have the same three column labels as in `wheel`. Once you've created that table, set `sum_bets` to your net gain in all 10 bets, assuming that you bet one dollar on red each time.

*Hint:* It may be helpful to print out `ten_bets` after you create it! 

```python
ten_bets = ...
sum_bets = ...(ten_bets.column('Winnings: Red'))
sum_bets
```

```python
grader.check("q1_3")
```

Run the cells above a few times to see how much money you would make if you made 10 one-dollar bets on red. Making a negative amount of money doesn't feel good, but it is a reality in gambling. Casinos are a business, and they make money when gamblers lose.

**Question 4.** Let's see what would happen if you made more bets. Define a function `net_gain_red` that takes the number of bets and returns the net gain in that number of one-dollar bets on red.

*Hint:* You should use your `wheel` table within your function and run the simulation like the in previous question.


```python
def net_gain_red(...):
    ...
```

```python
grader.check("q1_4")
```

Run the cell below a few times to make sure that the results are similar to those you observed in the previous exercise.

```python
net_gain_red(10)
```

**Question 5.** Complete the cell below to simulate the net gain in 200 one-dollar bets on red, repeating the process 10,000 times. After the cell is run, `simulated_gains_red` should be an array with 10,000 entries, each of which is the net gain in 200 one-dollar bets on red.

*Hint:* Think about which computational tool might be helpful for simulating a process multiple times. Lab 5 might be a good resource to look at!

*Note:* This cell might take a few seconds to run.


```python
num_bets = ...
repetitions = ...

simulated_gains_red = ...
...

len(simulated_gains_red) # Do not change this line! Check that simulated_gains_red is length 10000.
```

```python
grader.check("q1_5")
```

Run the cell below to visualize the results of your simulation.

```python
gains = Table().with_columns('Net Gain on Red', simulated_gains_red)
gains.hist(bins = np.arange(-80, 41, 4))
```

**Question 6:** Using the histogram above, decide whether the following statement is true or false:

>If you make 200 one-dollar bets on red, your chance of losing money is more than 50%.

Assign `loss_more_than_50` to either `True` or `False` depending on your answer to the question. **(2 points)**


```python
loss_more_than_50 = ...
```

```python
grader.check("q1_6")
```

### Betting on a Split ###
If betting on red doesn't seem like a good idea, maybe a gambler might want to try a different bet. A bet on a *split* is a bet on two consecutive numbers such as 5 and 6. This bets pays 17 to 1. That means if you place a one-dollar bet on the split 5 and 6, then:

- If the winning pocket is either 5 or 6, your gain is 17 dollars.
- If any other pocket wins, you lose your dollar, so your gain is -1 dollars.

**Question 7.** Define a function `dollar_bet_on_split` that takes a pocket number and returns your gain in dollars if that pocket won and you had bet one dollar on the 5-6 split.

*Hint:* Remember that the pockets are represented as strings.


```python
def dollar_bet_on_split(...):
    ...
```

```python
grader.check("q1_7")
```

Run the cell below to check that your function is doing what it should.

```python
print(dollar_bet_on_split('5'))
print(dollar_bet_on_split('6'))
print(dollar_bet_on_split('00'))
print(dollar_bet_on_split('23'))
```

**Question 8.** Add a column `Winnings: Split` to the `wheel` table. For each pocket, the column should contain your gain in dollars if that pocket won and you had bet one dollar on the 5-6 split.


```python
split_winnings = ...
wheel = ...
wheel.show(5) # Do not change this line.
```

```python
grader.check("q1_8")
```

**Question 9.** Simulate the net gain in 200 one-dollar bets on the 5-6 split, repeating the process 10,000 times and saving your gains in the array `simulated_gains_split`.

*Hint:* Your code in Questions 4 and 5 may be helpful here!


```python
simulated_gains_split = ...

...

# Do not change the two lines below
gains = gains.with_columns('Net Gain on Split', simulated_gains_split)
gains.hist(bins = np.arange(-200, 150, 20))
```

```python
grader.check("q1_9")
```

**Question 10.** Look carefully at the visualization above, and assign `histogram_statements` to an array of the numbers of each statement below that can be correctly inferred from the overlaid histogram.

1. If you bet one dollar 200 times on a split, your chance of losing money is more than 50%.
2. If you bet one dollar 200 times in roulette, your chance of making more than 50 dollars is greater if you bet on a split each time than if you bet on red each time.
3. If you bet one dollar 200 times in roulette, your chance of losing more than 50 dollars is greater if you bet on a split each time than if you bet on red each time.

*Hint:* We've already seen one of these statements in a prior question.


```python
histogram_statements = ...
```

```python
grader.check("q1_10")
```

If this exercise has put you off playing roulette, it has done its job. If you are still curious about other bets, [here](https://en.wikipedia.org/wiki/Roulette#Bet_odds_table) they all are, and [here](https://en.wikipedia.org/wiki/Roulette#House_edge) is the bad news. The house – that is, the casino – always has an edge over the gambler.

<!-- END QUESTION -->

## 2. Assessing Jade's Models

Before you begin, [Section 10.4](https://inferentialthinking.com/chapters/10/4/Random_Sampling_in_Python.html) of the textbook is a useful reference for this part.

### Games with Jade
Our friend Jade comes over and asks us to play a game with her. The game works like this: 

> We will draw randomly with replacement from a simplified 13 card deck with 4 face cards (A, J, Q, K), and 9 numbered cards (2, 3, 4, 5, 6, 7, 8, 9, 10). If we draw cards with replacement 13 times, and if the number of face cards is greater than or equal to 4, we lose.
> 
> Otherwise, Jade loses.

We play the game once and we lose, observing 8 total face cards. We are angry and accuse Jade of cheating! Jade is adamant, however, that the deck is fair.

Jade's model claims that there is an equal chance of getting any of the cards (A, 2, 3, 4, 5, 6, 7, 8, 9, 10, J, Q, K), but we do not believe her. We believe that the deck is clearly rigged, with face cards (A, J, Q, K) being more likely than the numbered cards (2, 3, 4, 5, 6, 7, 8, 9, 10).

**Question 1.** Assign `deck_model_probabilities` to a two-item array containing the chance of drawing a face card as the first element, and the chance of drawing a numbered card as the second element under Jade's model. Since we're working with probabilities, make sure your values are between 0 and 1.


```python
deck_model_probabilities = ...
deck_model_probabilities
```

```python
grader.check("q4_1")
```

**Question 2.** We believe Jade's model is incorrect. In particular, we believe there to be a *larger* chance of getting a face card. Which of the following statistics can we use during our simulation to test between the model and our alternative? Assign `statistic_choice` to the correct answer.

1. The distance (absolute value) between the actual number of face cards in 13 draws and 4, the expected number of face cards in 13 draws
2. The expected number of face cards in 13 draws
3. The number of face cards we get in 13 draws


```python
statistic_choice = ...
statistic_choice
```

```python
grader.check("q4_2")
```

**Question 3.** Define the function `deck_simulation_and_statistic`, which, given a sample size and an array of model proportions (like the one you created in Question 1), returns the **number of face cards** in one simulation of drawing cards under the model specified in `model_proportions`.

*Hint:* Think about how you can use the function `sample_proportions`. 


```python
def deck_simulation_and_statistic(sample_size, model_proportions):
    ...

deck_simulation_and_statistic(13, deck_model_probabilities)
```

```python
grader.check("q4_3")
```

**Question 4.** Use your function from above to simulate the drawing of 13 cards 5000 times under the proportions that you specified in Question 1. Keep track of all of your statistics in `deck_statistics`.


```python
repetitions = 5000 
...

deck_statistics
```

```python
grader.check("q4_4")
```

Let’s take a look at the distribution of simulated statistics.

```python
# Draw a distribution of statistics 
Table().with_column('Deck Statistics', deck_statistics).hist()
```

<!-- BEGIN QUESTION -->

**Question 5.** Given your observed value, do you believe that Jade's model is reasonable, or is our alternative (that our deck is rigged) more likely? Explain your answer using the histogram of statistics simulated using Jade's model (produced above).


_Type your answer here, replacing this text._

Tests for q1_1 (points: [2])
```
>>> dollar_bet_on_red('red') == 1
True
```


Tests for q1_10 (points: [0, 0, 0])
```
>>> # Make sure histogram_statements is an array.
>>> import numpy as np
>>> type(histogram_statements) == np.ndarray
True
```
```
>>> # Make sure you have set histogram_statements to an array with at least 1 number
>>> 1 <= histogram_statements.item(0) <= 3
True
```
```
>>> # make sure all elements are between 1 and 3
>>> all(histogram_statements <= 3) and all(histogram_statements >= 1)
True
```


Tests for q1_2 (points: [0])
```
>>> # Make sure your column names are correct
>>> wheel.labels[2] == "Winnings: Red"
True
```


Tests for q1_3 (points: [0, 0])
```
>>> # If this test isn't passing, try running the cells from the top.
>>> set(["Pocket", "Color", "Winnings: Red"]) == set(ten_bets.labels)
True
```
```
>>> # Checks that ten_bets doesn't have the same number of rows as wheel
>>> ten_bets.num_rows != wheel.num_rows
True
```


Tests for q1_4 (points: [0])
```
>>> net_gain_red(10000) != net_gain_red(10000)
True
```


Tests for q1_5 (points: [0])
```
>>> # Make sure simulated_gains_red is an array.
>>> import numpy as np
>>> type(simulated_gains_red) == np.ndarray
True
```


Tests for q1_6 (points: [0])
```
>>> type(loss_more_than_50) == bool
True
```


Tests for q1_7 (points: [2])
```
>>> dollar_bet_on_split('5') == 17
True
```


Tests for q1_8 (points: [0])
```
>>> # Make sure your column names are correct
>>> wheel.labels[3] == "Winnings: Split"
True
```


Tests for q1_9 (points: [0])
```
>>> # Make sure simulated_gains_split is an array.
>>> import numpy as np
>>> type(simulated_gains_split) == np.ndarray
True
```


Tests for q2_1 (points: [0])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < first_three_black < 1
True
```


Tests for q2_2 (points: [0])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < no_green < 1
True
```


Tests for q2_3 (points: [0])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < at_least_one_green < 1
True
```


Tests for q2_4 (points: [0])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < lone_winners < 1
True
```


Tests for q2_5 (points: [0, 0])
```
>>> # Make sure you are setting toys_option to an int
>>> type(toys_option) == int
True
```
```
>>> # toys_option should be assigned to either 1, 2 or 3.
>>> 1 <= toys_option <= 3
True
```


Tests for q2_6 (points: [0, 0])
```
>>> # Make sure you are setting lottery_option to an int
>>> type(lottery_option) == int
True
```
```
>>> # toys_option should be assigned to either 1, 2 or 3.
>>> 1 <= lottery_option <= 3
True
```


Tests for q2_7 (points: [0, 0])
```
>>> # Make sure you are setting coin_option to an int
>>> type(coin_option) == int
True
```
```
>>> 1 <= coin_option <= 3
True
```


Tests for q4_1 (points: [0, 0])
```
>>> # The array should have length 2
>>> len(deck_model_probabilities) == 2
True
```
```
>>> # The elements in the array should add up to 1.
>>> sum(deck_model_probabilities) == 1
True
```


Tests for q4_2 (points: [0, 0])
```
>>> # Make sure you are setting statistic_choice to an int
>>> type(statistic_choice) == int
True
```
```
>>> # statistic_choice should be assigned to either 1, 2 or 3.
>>> 1 <= statistic_choice <= 3
True
```


Tests for q4_3 (points: [0, 5])
```
>>> # Make sure to return the number of face cards, not a proportion
>>> num_face = deck_simulation_and_statistic(13, deck_model_probabilities)
>>> num_face % 1 == 0
True
```
```
>>> # The statistic should be between 0 and 13 face cards for
>>> # a sample size of 13
>>> num_face = deck_simulation_and_statistic(13, deck_model_probabilities)
>>> 0 <= num_face <= 13
True
```


Tests for q4_4 (points: [2, 3])
```
>>> # There should be exactly as many elements in deck_statistics
>>> # as the number 'repetitions'
>>> len(deck_statistics) == repetitions
True
```
```
>>> # Each element of deck_statistics should be between 0
>>> # and 13 inclusive
>>> all([0 <= k <= 13 for k in deck_statistics])
True
```
