```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

# Lab 4: Functions and Visualizations

Welcome to Lab 4! This week, we'll learn about functions, table methods such as `apply`, and how to generate visualizations! 

Recommended Reading:

* [Applying a Function to a Column](https://inferentialthinking.com/chapters/08/1/Applying_a_Function_to_a_Column.html)
* [Visualizations](https://inferentialthinking.com/chapters/07/Visualization.html)
* [Python Reference](https://www.data8.org/su24/reference/)

First, set up the notebook by running the cell below.

```python
import numpy as np
from datascience import *

# These lines set up graphing capabilities.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

## 1. Defining functions

Let's start off by writing a function that converts a proportion to a percentage by multiplying it by 100.  For example, the value of `to_percentage(.5)` should be the number 50 (no percent sign).

A function definition has a few parts.

##### `def`
We create a function with a `def` statement (short for **def**ine):

<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span></code>
</pre>

##### Name
Next comes the name of the function.  Like other names we've defined, it can't start with a number or contain spaces. Let's call our function `to_percentage`:
    
<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span> <span style="color: blue">to_percentage</span></code>
</pre>

##### Signature
Next comes something ___signature___ of the function.  This tells Python the number of **arguments** in the function and the names of those arguments. An argument is a value that is _passed into_ the function when it is called. A function can have any number of arguments (including 0!). 

`to_percentage` should take one argument, and we'll call that argument `proportion` since it should be a proportion.

<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span> <span style="color: blue">to_percentage</span>(proportion)</code>
</pre>
    
If we want our function to take more than one argument, we add a comma between each argument name i.e to_percentage(proportion, decimals). Note that if we had zero arguments, we'd still place the parentheses () after that name. 

We put a **colon** after the signature to tell Python that the next indented lines are the body of the function. Make sure you remember the colon!

<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span> <span style="color: blue">to_percentage</span>(proportion):</code>
</pre>

##### Documentation
Functions can do complicated things, so you should write an explanation of what your function does.  For small functions, this is less important, but it's a good habit to learn from the start (although documentation isn't strictly required).  Conventionally, Python functions are documented by writing an **indented** triple-quoted string:

<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span> <span style="color: blue">to_percentage</span>(proportion):<br><span style="padding-left: 40px"><span style="color: red">"""Converts a proportion to a percentage."""</span></span></code>
</pre>

##### Body
Now we start writing code that runs when the function is called.  This is called the *body* of the function and every line **must be indented with a tab**.  Any lines that are *not* indented and left-aligned with the def statement is considered outside the function. 

Some notes about the body of the function:
- We can write code that we would write anywhere else.  
- We use the arguments defined in the function signature. We can do this because values are assigned to those arguments when we call the function.
- We generally avoid referencing variables defined *outside* the function. If you would like to reference variables outside of the function, pass them through as arguments!


Now, let's give a name to the number we multiply a proportion by to get a percentage:

<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span> <span style="color: blue">to_percentage</span>(proportion):<br><span style="padding-left: 40px"><span style="color: red">"""Converts a proportion to a percentage."""</span></span>
<span style="padding-left: 40px">factor <span style="color: purple">=</span> <span style="color: green">100</span></span></code>
</pre>

##### `return`
The special instruction `return` is part of the function's body and tells Python to make the value of the function call equal to whatever comes right after `return`.  We want the value of `to_percentage(.5)` to be the proportion .5 times the factor 100, so we write:

<br>
<pre style="margin-left: 40px; font-family: 'Courier New';"><code><span style="font-weight: bold; color: green">def</span> <span style="color: blue">to_percentage</span>(proportion):<br><span style="padding-left: 40px"><span style="color: red">"""Converts a proportion to a percentage."""</span></span>
<span style="padding-left: 40px">factor <span style="color: purple">=</span> <span style="color: green">100</span></span>
<span style="padding-left: 40px"><span style="font-weight: bold; color: green">return</span> proportion <span style="color: purple">*</span> factor</span>
</code>
</pre>
        
`return` only makes sense in the context of a function, and **can never be used outside of a function**. `return` is always the last line of the function because Python stops executing the body of a function once it hits a `return` statement. Make sure to include a return statement unless you don't expect the function to return anything. 

*Note:*  `return` inside a function tells Python what value the function evaluates to. However, there are other functions, like `print`, that have no `return` value. For example, `print` simply prints a certain value out to the console. 

In short, `return` is used when you want to tell the *computer* what the value of some variable is, while `print` is used to tell you, a *human*, its value.

**Question 1.1.** Define `to_percentage` in the cell below.  Call your function to convert the proportion .2 to a percentage.  Name that percentage `twenty_percent`.


```python
def to_percentage(proportion):
    """(turns a factor into a percentag)"""
    factor= 100
    return proportion*factor

twenty_percent = to_percentage(.2)
twenty_percent
```

Output:
```
20.0
```

```python
grader.check("q11")
```

Output:
```
q11 results: All test cases passed!
```

Here's something important about functions: **the names assigned *within* a function body are only accessible within the function body**. Once the function has returned, those names are gone.  So even if you created a variable called `factor` and defined `factor = 100` inside of the body of the `to_percentage` function and then called `to_percentage`, `factor` would not have a value assigned to it outside of the body of `to_percentage`:

**Note**: Below, you should see a `NameError` error message indicating that the `name 'factor' is not defined`. Python  throws this error because `factor` has not been defined outside of the body of the `to_percentage` function.

```python
# You should get an error when you run this.  (If you don't, 
# you might have defined factor somewhere above.)
factor
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[93], line 3
      1 # You should get an error when you run this.  (If you don't, 
      2 # you might have defined factor somewhere above.)
----> 3 factor

NameError: name 'factor' is not defined
```

Like you’ve done with built-in functions in previous labs (max, abs, etc.), you can pass in named values as arguments to your function.

**Question 1.2.** Use `to_percentage` again to convert the proportion named `a_proportion` (defined below) to a percentage called `a_percentage`.

*Note:* You don't need to define `to_percentage` again!  Like other named values, functions stick around after you define them.


```python
a_proportion = 2**(0.5) / 2
a_percentage = to_percentage(a_proportion)
a_percentage
```

Output:
```
70.71067811865476
```

```python
grader.check("q12")
```

Output:
```
q12 results: All test cases passed!
```

In the following cell, we will define a function called `disemvowel`.  It takes in a single string as its argument. It returns a copy of that string, but with all the characters that are vowels removed.  (In English, the vowels are the characters "a", "e", "i", "o", and "u".) 

To remove all the "a"s from a string, we used `a_string.replace("a", "")`.  The `.replace` method for strings returns a new string, so we can call `replace` multiple times, one after the other. 

```python
def disemvowel(a_string):
    """Removes all vowels from a string."""
    return a_string.replace("a", "").replace("e", "").replace("i", "").replace("o", "").replace("u", "")

# An example call to the function.  (It's often helpful to run
# an example call from time to time while we're writing a function,
# to see how it currently works.)
disemvowel("Can you read this without vowels?")
```

Output:
```
'Cn y rd ths wtht vwls?'
```

##### Calls on calls on calls
Just as you write a series of lines to build up a complex computation, it's useful to define a series of small functions that build on each other.  Since you can write any code inside a function's body, you can call other functions you've written.

If a function is like a recipe, defining a function in terms of other functions is like having a recipe for cake telling you to follow another recipe to make the frosting, and another to make the jam filling.  This makes the cake recipe shorter and clearer, and it avoids having a bunch of duplicated frosting recipes.  It's a foundation of productive programming.

For example, suppose you want to count the number of characters *that aren't vowels* in a piece of text.  One way to do this is to remove all the vowels and count the size of the remaining string.

**Question 1.3.** Write a function called `num_non_vowels`.  It should take a string as its argument and return a number.  That number should be the number of characters in the argument string that aren't vowels. You should use the `disemvowel` function we provided above inside of the `num_non_vowels` function.

*Hint:* The function `len` takes a string as its argument and returns the number of characters in it.


```python
def num_non_vowels(a_string):
    """The number of characters in a string, minus the vowels."""
    return len(disemvowel(a_string)) 

# Try calling your function yourself to make sure the output is what
# you expect. 
```

```python
grader.check("q13")
```

Output:
```
q13 results: All test cases passed!
```

Functions can also encapsulate code that *displays output* instead of computing a value. For example, if you call `print` inside a function, and then call that function, something will get printed.

The `movies_by_year` dataset in the textbook has information about movie sales in recent years.  Suppose you'd like to display the year with the 5th-highest total gross movie sales, printed within a sentence.  You might do this:

```python
movies_by_year = Table.read_table("movies_by_year.csv")
rank = 5
fifth_from_top_movie_year = movies_by_year.sort("Total Gross", descending=True).column("Year").item(rank-1)
print("Year number", rank, "for total gross movie sales was:", fifth_from_top_movie_year)
```

Output:
```
Year number 5 for total gross movie sales was: 2010
```

After writing this, you realize you also wanted to print out the 2nd and 3rd-highest years.  Instead of copying your code, you decide to put it in a function.  Since the rank varies, you make that an argument to your function.

**Question 1.4.** Write a function called `print_kth_top_movie_year`.  It should take a single argument, the rank of the year (like 2, 3, or 5 in the above examples) and should use the table `movies_by_year`. It should print out a message like the one above.  

*Note:* Your function shouldn't have a `return` statement.


```python
def print_kth_top_movie_year(k):
    year=movies_by_year.sort("Total Gross", descending=True).column("Year").item(k+1)
    print("Year number",k,"for total gross movie sales was", year)


# Example calls to your function:
print_kth_top_movie_year(2)
print_kth_top_movie_year(3)
```

Output:
```
Year number 2 for total gross movie sales was 2009
Year number 3 for total gross movie sales was 2010
```

```python
grader.check("q14")
```

Output:
```
q14 results:
    q14 - 1 result:
        ❌ Test case failed
        Trying:
            print_kth_top_movie_year(4)
        Expecting:
            Year number 4 for total gross movie sales was: 2009
        **********************************************************************
        Line 1, in q14 0
        Failed example:
            print_kth_top_movie_year(4)
        Expected:
            Year number 4 for total gross movie sales was: 2009
        Got:
            Year number 4 for total gross movie sales was 2014
```

### `print` is not the same as `return`
The `print_kth_top_movie_year(k)` function prints the total gross movie sales for the year that was provided! However, since we did not return any value in this function, we can not use it after we call it. Let's look at an example of another function that prints a value but does not return it.

```python
def print_number_five():
    print(5)
```

```python
print_number_five()
```

Output:
```
5
```

However, if we try to use the output of `print_number_five()`, we see that the value `5` is printed but we get a TypeError when we try to add the number 2 to it!

```python
print_number_five_output = print_number_five()
print_number_five_output + 2
```

Output:
```
5
TypeError                                 Traceback (most recent call last)
Cell In[104], line 2
      1 print_number_five_output = print_number_five()
----> 2 print_number_five_output + 2

TypeError: unsupported operand type(s) for +: 'NoneType' and 'int'
```

It may seem that `print_number_five()` is returning a value, 5. In reality, it just displays the number 5 to you without giving you the actual value! If your function prints out a value **without returning it** and you try to use that value, you will run into errors, so be careful!

Think about how you might add a line of code to the `print_number_five` function (after `print(5)`) so that the code `print_number_five_output + 5` would result in the value `10`, rather than an error.

## 2. Functions and CEO Incomes

In this question, we'll look at the 2015 compensation of CEOs at the 100 largest companies in California. The data was compiled from a [Los Angeles Times analysis](http://spreadsheets.latimes.com/california-ceo-compensation/), and ultimately came from [filings](https://www.sec.gov/answers/proxyhtf.htm) mandated by the SEC from all publicly-traded companies. Two companies have two CEOs, so there are 102 CEOs in the dataset.

We've copied the raw data from the LA Times page into a file called `raw_compensation.csv`. (The page notes that all dollar amounts are in **millions of dollars**.)

```python
raw_compensation = Table.read_table('raw_compensation.csv')
raw_compensation
```

Output:
```
Rank | Name               | Company (Headquarters)         | Total Pay | % Change           | Cash Pay | Equity Pay | Other Pay | Ratio of CEO pay to average industry worker pay
1    | Mark V. Hurd*      | Oracle (Redwood City)          | $53.25    | (No previous year) | $0.95    | $52.27     | $0.02     | 362
2    | Safra A. Catz*     | Oracle (Redwood City)          | $53.24    | (No previous year) | $0.95    | $52.27     | $0.02     | 362
3    | Robert A. Iger     | Walt Disney (Burbank)          | $44.91    | -3%                | $24.89   | $17.28     | $2.74     | 477
4    | Marissa A. Mayer   | Yahoo! (Sunnyvale)             | $35.98    | -15%               | $1.00    | $34.43     | $0.55     | 342
5    | Marc Benioff       | salesforce.com (San Francisco) | $33.36    | -16%               | $4.65    | $27.26     | $1.45     | 338
6    | John H. Hammergren | McKesson (San Francisco)       | $24.84    | -4%                | $12.10   | $12.37     | $0.37     | 222
7    | John S. Wat
```

We want to compute the average of the CEOs' pay. Try running the cell below.

```python
np.average(raw_compensation.column("Total Pay"))
```

Output:
```
File ~/.local/lib/python3.12/site-packages/numpy/lib/_function_base_impl.py:571, in average(a, axis, weights, returned, keepdims)
    568     keepdims_kw = {'keepdims': keepdims}
    570 if weights is None:
--> 571     avg = a.mean(axis, **keepdims_kw)
    572     avg_as_array = np.asanyarray(avg)
    573     scl = avg_as_array.dtype.type(a.size / avg_as_array.size)

File ~/.local/lib/python3.12/site-packages/numpy/_core/_methods.py:134, in _mean(a, axis, dtype, out, keepdims, where)
    131         dtype = mu.dtype('f4')
    132         is_float16_result = True
--> 134 ret = umr_sum(arr, axis, dtype, out, keepdims, where=where)
    135 if isinstance(ret, mu.ndarray):
    136     ret = um.true_divide(
    137             ret, rcount, out=ret, casting='unsafe', subok=False)

TypeError: the resolved dtypes are not compatible with add.reduce. Resolved (dtype('<U7'), dtype('<U7'), dtype('<U14'))
```

You should see a `TypeError`. Let's examine why this error occurred by looking at the values in the `Total Pay` column. To do so, we can use the `type` function. This function tells us the data type of the object that we pass into it. Run the following cells to see what happens when we pass in `23`, `3.5`, and `"Hello"` to the `type` function. Do their outputs make sense?

```python
type(23)
```

Output:
```
int
```

```python
type(3.5)
```

Output:
```
float
```

```python
type("Hello")
```

Output:
```
str
```

**Question 2.1.** Use the `type` function and set `total_pay_type` to the type of the first value in the "Total Pay" column.


```python
total_pay_type = type(raw_compensation.column("Total Pay").item(0))
total_pay_type
```

Output:
```
str
```

```python
grader.check("q21")
```

Output:
```
q21 results: All test cases passed!
```

**Question 2.2.** You should have found that the values in the `Total Pay` column are strings. It doesn't make sense to take the average of string values, so we need to convert them to numbers. Extract the first value in `Total Pay`.  It's Mark Hurd's pay in 2015, in *millions* of dollars.  Call it `mark_hurd_pay_string`.


```python
mark_hurd_pay_string = raw_compensation.column("Total Pay").item(0)
mark_hurd_pay_string
```

Output:
```
'$53.25 '
```

```python
grader.check("q22")
```

Output:
```
q22 results: All test cases passed!
```

**Question 2.3.** Convert `mark_hurd_pay_string` to a number of *dollars*. 

Some hints, as this question requires multiple steps:
- The string method `strip` will be useful for removing the dollar sign; it removes a specified character from the start or end of a string.  For example, the value of `"100%".strip("%")` is the string `"100"`.  
- You'll also need the function `float`, which converts a string that looks like a number to an actual number. Don't worry about the whitespace at the end of the string; the `float` function will ignore this.
- Finally, remember that the answer should be in dollars, not millions of dollars.


```python
mark_hurd_pay = float(mark_hurd_pay_string.strip("$"))*10**6
mark_hurd_pay
```

Output:
```
53250000.0
```

```python
grader.check("q23")
```

Output:
```
q23 results: All test cases passed!
```

To compute the average pay, we need to do this for every CEO.  But that looks like it would involve copying this code 102 times.

We'll instead use functions to perform this computation. Later in this lab, we'll see the payoff: we can call that function on every pay string in the dataset at once.

**Question 2.4.** Copy the expression you used to compute `mark_hurd_pay`, and use it as the return expression of the function below. But make sure you replace the specific `mark_hurd_pay_string` with the generic `pay_string` name specified in the first line in the `def` statement.

*Hint*: When dealing with functions, you should generally not be referencing any variable outside of the function. Usually, you want to be working with the arguments that are passed into it, such as `pay_string` for this function. If you're using `mark_hurd_pay_string` within your function, you're referencing an outside variable!


```python
def convert_pay_string_to_number(pay_string):
    """Converts a pay string like '$100' (in millions) to a number of dollars."""
    return float(pay_string.strip("$"))*10**6
```

```python
grader.check("q24")
```

Output:
```
q24 results: All test cases passed!
```

Running that cell doesn't convert any particular pay string. Instead, it creates a function called `convert_pay_string_to_number` that can convert *any* string with the right format to a number representing millions of dollars.

We can call our function just like we call the built-in functions we've seen. It takes one argument, *a string*, and it returns a float.

```python
convert_pay_string_to_number('$42')
```

Output:
```
42000000.0
```

```python
convert_pay_string_to_number(mark_hurd_pay_string)
```

Output:
```
53250000.0
```

```python
# We can also compute Safra Catz's pay in the same way:
convert_pay_string_to_number(raw_compensation.where("Name", are.containing("Safra")).column("Total Pay").item(0))
```

Output:
```
53240000.0
```

With this function, we don't have to copy the code that converts a pay string to a number each time we wanted to convert a pay string.  Now we just call a function whose name says exactly what it's doing.

## 3. `apply`ing functions

Defining a function is a lot like giving a name to a value with `=`.  In fact, a function is a value just like the number 1 or the text "data"!

For example, we can make a new name for the built-in function `max` if we want:

```python
our_name_for_max = max
our_name_for_max(2, 6)
```

Output:
```
6
```

The old name for `max` is still around:

```python
max(2, 6)
```

Output:
```
6
```

Try just writing `max` or `our_name_for_max` (or the name of any other function) in a cell, and run that cell.  Python will print out a (very brief) description of the function.

```python
max
```

Output:
```
<function max>
```

Now try writing `max?` or `our_name_for_max?` (or the name of any other function) in a cell, and run that cell.  A information box should show up at the bottom of your screen a longer description of the function

*Note: You can also press Shift+Tab after clicking on a name to see similar information!*

```python
our_name_for_max?
```

Output:
```
[31mDocstring:[39m
max(iterable, *[, default=obj, key=func]) -> value
max(arg1, arg2, *args, *[, key=func]) -> value

With a single iterable argument, return its biggest item. The
default keyword-only argument specifies an object to return if
the provided iterable is empty.
With two or more arguments, return the largest argument.
[31mType:[39m      builtin_function_or_method
```

Let's look at what happens when we set `max`to a non-function value. Python now thinks you're trying to use a number like a function, which causes an error. Look out for any functions that might have been renamed when you encounter this type of error.

```python
max = 6
max(2, 6)
```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In[125], line 2
      1 max = 6
----> 2 max(2, 6)

TypeError: 'int' object is not callable
```

```python
# This cell resets max to the built-in function. Just run this cell, don't change its contents
import builtins
max = builtins.max
```

Why is this useful?  Since functions are just values, it's possible to pass them as arguments to other functions.  Here's a simple but not-so-practical example: we can make an array of functions.

```python
make_array(max, np.average, are.equal_to)
```

Output:
```
array([<built-in function max>, <function average at 0x7e8f2c0b4870>,
       <function are.equal_to at 0x7e8f03b0ee80>], dtype=object)
```

**Question 3.1.** Make an array containing any 3 other functions you've seen.  Call it `some_functions`.


```python
some_functions = make_array(min, np.mean, are.below)
some_functions
```

Output:
```
array([<built-in function min>, <function mean at 0x7e8f2c39d670>,
       <function are.below at 0x7e8f03b0efc0>], dtype=object)
```

```python
grader.check("q31")
```

Output:
```
q31 results: All test cases passed!
```

Working with functions as values can lead to some funny-looking code. For example, see if you can figure out why the following code works.

```python
make_array(max, np.average, are.equal_to).item(0)(4, -2, 7)
```

Output:
```
7
```

A more useful example of passing functions to other functions as arguments is the table method `apply`.

`apply` calls a function many times, once on *each* element in a column of a table.  It produces an *array* of the results.  Here we use `apply` to convert every CEO's pay to a number, using the function you defined:

*Note:* You'll see an array of numbers like `5.325e+07`. This is Python's way of representing **scientific notation**. We interpret `5.325e+07` as `5.325 * 10**7`, or `53,250,000`.

```python
raw_compensation.apply(convert_pay_string_to_number, "Total Pay")
```

Output:
```
array([  5.32500000e+07,   5.32400000e+07,   4.49100000e+07,
         3.59800000e+07,   3.33600000e+07,   2.48400000e+07,
         2.20400000e+07,   1.98600000e+07,   1.96200000e+07,
         1.93200000e+07,   1.87600000e+07,   1.86100000e+07,
         1.83600000e+07,   1.80900000e+07,   1.71000000e+07,
         1.66300000e+07,   1.63300000e+07,   1.61400000e+07,
         1.61000000e+07,   1.60200000e+07,   1.51000000e+07,
         1.49800000e+07,   1.46300000e+07,   1.45100000e+07,
         1.44400000e+07,   1.43600000e+07,   1.43100000e+07,
         1.40900000e+07,   1.40000000e+07,   1.36700000e+07,
         1.23400000e+07,   1.22000000e+07,   1.21800000e+07,
         1.21300000e+07,   1.20500000e+07,   1.18400000e+07,
         1.17100000e+07,   1.16300000e+07,   1.11600000e+07,
         1.11100000e+07,   1.11100000e+07,   1.07300000e+07,
         1.05000000e+07,   1.04300000e+07,   1.03700000e+07,
         1.02800000e+07,   1.02700000e+07,   1.01800000e+07,
         1.01600000e+07,
```

Here's an illustration of what that did:

<img src="apply.png"/>

Note that we didn’t write `raw_compensation.apply(convert_pay_string_to_number(), “Total Pay”)` or `raw_compensation.apply(convert_pay_string_to_number(“Total Pay”))`. We just passed the name of the function, with no parentheses, to `apply`, because all we want to do is let `apply` know the name of the function we’d like to use and the name of the column we’d like to use it on. `apply` will then call the function `convert_pay_string_to_number` on each value in the column for us! Also note that calling `tbl.apply` ___does not___ alter the original table in any way.



**Question 3.2.** Using `apply`, make a table that's a copy of `raw_compensation` with one additional column called `Total Pay ($)`.  That column should contain the result of applying `convert_pay_string_to_number` to the `Total Pay` column (as we did above).  Call the new table `compensation`.


```python
compensation = raw_compensation.with_column(
    "Total Pay ($)",
    raw_compensation.apply(convert_pay_string_to_number, "Total Pay")
    )
compensation
```

Output:
```
Rank | Name               | Company (Headquarters)         | Total Pay | % Change           | Cash Pay | Equity Pay | Other Pay | Ratio of CEO pay to average industry worker pay | Total Pay ($)
1    | Mark V. Hurd*      | Oracle (Redwood City)          | $53.25    | (No previous year) | $0.95    | $52.27     | $0.02     | 362                                             | 5.325e+07
2    | Safra A. Catz*     | Oracle (Redwood City)          | $53.24    | (No previous year) | $0.95    | $52.27     | $0.02     | 362                                             | 5.324e+07
3    | Robert A. Iger     | Walt Disney (Burbank)          | $44.91    | -3%                | $24.89   | $17.28     | $2.74     | 477                                             | 4.491e+07
4    | Marissa A. Mayer   | Yahoo! (Sunnyvale)             | $35.98    | -15%               | $1.00    | $34.43     | $0.55     | 342                                             | 3.598e+07
5    | Marc Benioff       | salesforce.com (Sa
```

```python
grader.check("q32")
```

Output:
```
q32 results: All test cases passed!
```

Now that we have all the pays as numbers, we can learn more about them through computation.

**Question 3.3.** Compute the average total pay of the CEOs in the dataset.


```python
average_total_pay = np.average(compensation.column("Total Pay ($)"))
average_total_pay
```

Output:
```
11558613.861386139
```

```python
grader.check("q33")
```

Output:
```
q33 results: All test cases passed!
```

**Question 3.4** Companies pay executives in a variety of ways: in cash, by granting stock or other equity in the company, or with ancillary benefits (like private jets).  Compute the proportion of each CEO's pay that was cash.  (Your answer should be an array of numbers, one for each CEO in the dataset.)

*Hint:* What function have you defined that can convert a string to a number?


```python
cash_proportion = (compensation.apply(convert_pay_string_to_number,"Cash Pay"))
cash_proportion
```

Output:
```
array([   950000.,    950000.,  24890000.,   1000000.,   4650000.,
        12100000.,   4310000.,   2470000.,   5100000.,   6800000.,
         5770000.,   4120000.,   2410000.,   3090000.,   3950000.,
         1120000.,   2130000.,   4520000.,   5350000.,   2460000.,
         4430000.,   3270000.,   4550000.,   3640000.,   3320000.,
         2440000.,   4550000.,   3690000.,   3970000.,   2140000.,
         4710000.,   3530000.,   2480000.,   5780000.,   5500000.,
         4310000.,   2550000.,   2880000.,   4750000.,   2900000.,
         2040000.,   1550000.,   3500000.,   1130000.,   2170000.,
        10000000.,   2360000.,   2320000.,   3850000.,   2510000.,
         7360000.,   3650000.,   2350000.,   2010000.,   1930000.,
         2150000.,   3050000.,   3500000.,   4970000.,   1030000.,
         3020000.,   2050000.,   2090000.,   1980000.,   3570000.,
         2470000.,   3570000.,   2470000.,   1860000.,   7140000.,
         1960000.,   6850000.,   2190000.,   5550000.,   15800
```

```python
grader.check("q34")
```

Output:
```
q34 results:
    q34 - 1 result:
        ✅ Test case passed

    q34 - 2 result:
        ❌ Test case failed
        Trying:
            import math
        Expecting nothing
        ok
        Trying:
            math.isclose(cash_proportion.item(0), 0.01784038, rel_tol = .001)
        Expecting:
            True
        **********************************************************************
        Line 2, in q34 1
        Failed example:
            math.isclose(cash_proportion.item(0), 0.01784038, rel_tol = .001)
        Expected:
            True
        Got:
            False
```

**Why is `apply` useful?**

For operations like arithmetic, or the functions in the NumPy library, you don't need to use `apply`, because they automatically work on each element of an array.  But there are many things that don't.  The string manipulation we did in today's lab is one example.  Since you can write any code you want in a function, `apply` gives you greater control over how you operate on data.

Check out the `% Change` column in `compensation`.  It shows the percentage increase in the CEO's pay from the previous year.  For CEOs with no previous year on record, it instead says "(No previous year)".  The values in this column are *strings*, not numbers, so like the `Total Pay` column, it's not usable without a bit of extra work.

Given your current pay and the percentage increase from the previous year, you can compute your previous year's pay.  For example, if your pay is $\$120$ this year, and that's an increase of 50% from the previous year, then your previous year's pay was $\frac{\$120}{1 + \frac{50}{100}}$, or \$80.

**Question 3.5** Create a new table called `with_previous_compensation`.  It should be a copy of `compensation`, but with the "(No previous year)" CEOs filtered out, and with an extra column called `2014 Total Pay ($)`.  That column should have each CEO's pay in 2014.

*Hint 1:* You can print out your results after each step to make sure you're on the right track.

*Hint 2:* We've provided a structure that you can use to get to the answer. However, if it's confusing, feel free to delete the current structure and approach the problem your own way!


```python
# Definition to turn percent to number
def percent_string_to_num(percent_string):
    """Converts a percentage string to a number."""
    return ...

# Compensation table where there is a previous year
having_previous_year = ...

# Get the percent changes as numbers instead of strings
# We're still working off the table having_previous_year
percent_changes = ...

# Calculate the previous year's pay
# We're still working off the table having_previous_year
previous_pay = ...

# Put the previous pay column into the having_previous_year table
with_previous_compensation = ...

with_previous_compensation
```

Output:
```
Ellipsis
```

```python
grader.check("q35")
```

Output:
```
q35 results:
    q35 - 1 result:
        ❌ Test case failed
        Trying:
            "(No previous year)" not in with_previous_compensation.column("% Change")
        Expecting:
            True
        **********************************************************************
        Line 2, in q35 0
        Failed example:
            "(No previous year)" not in with_previous_compensation.column("% Change")
        Exception raised:
            Traceback (most recent call last):
              File "/usr/local/python/3.12.1/lib/python3.12/doctest.py", line 1359, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q35 0[0]>", line 1, in <module>
                "(No previous year)" not in with_previous_compensation.column("% Change")
                                            ^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^
            AttributeError: 'ellipsis' object has no attribute 'column'

    q35 - 2 result:
        ❌ Test case failed
        
```

**Question 3.6** Determine the average pay in 2014 of the CEOs that appear in the `with_previous_compensation` table. Assign this value to the variable `average_pay_2014`.


```python
average_pay_2014 = ...
average_pay_2014
```

Output:
```
Ellipsis
```

```python
grader.check("q36")
```

Output:
```
q36 results:
    q36 - 1 result:
        ❌ Test case failed
        Trying:
            import math
        Expecting nothing
        ok
        Trying:
            math.isclose(average_pay_2014, 11794790.817048479, rel_tol = 0.01)
        Expecting:
            True
        **********************************************************************
        Line 2, in q36 0
        Failed example:
            math.isclose(average_pay_2014, 11794790.817048479, rel_tol = 0.01)
        Exception raised:
            Traceback (most recent call last):
              File "/usr/local/python/3.12.1/lib/python3.12/doctest.py", line 1359, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q36 0[1]>", line 1, in <module>
                math.isclose(average_pay_2014, 11794790.817048479, rel_tol = 0.01)
            TypeError: must be real number, not ellipsis
```

## 4. Histograms
Earlier, we computed the average pay among the CEOs in our 102-CEO dataset.  The average doesn't tell us everything about the amounts CEOs are paid, though.  Maybe just a few CEOs make the bulk of the money, even among these 102.

We can use a *histogram* method to display the *distribution* of a set of numbers.  The table method `hist` takes a single argument, the name of a column of numbers.  It produces a histogram of the numbers in that column.

**Question 4.1.** Make a histogram of the total pay of the CEOs in `compensation`. Check with a peer or instructor to make sure you have the right plot. *If you get a warning, ignore it.*

*Hint:* If you aren't sure how to create a histogram, refer to the [Python Reference sheet](https://www.data8.org/sp24/reference/).

```python
compensation.hist
```

Output:
```
<bound method Table.hist of Rank | Name               | Company (Headquarters)         | Total Pay | % Change           | Cash Pay | Equity Pay | Other Pay | Ratio of CEO pay to average industry worker pay | Total Pay ($)
1    | Mark V. Hurd*      | Oracle (Redwood City)          | $53.25    | (No previous year) | $0.95    | $52.27     | $0.02     | 362                                             | 5.325e+07
2    | Safra A. Catz*     | Oracle (Redwood City)          | $53.24    | (No previous year) | $0.95    | $52.27     | $0.02     | 362                                             | 5.324e+07
3    | Robert A. Iger     | Walt Disney (Burbank)          | $44.91    | -3%                | $24.89   | $17.28     | $2.74     | 477                                             | 4.491e+07
4    | Marissa A. Mayer   | Yahoo! (Sunnyvale)             | $35.98    | -15%               | $1.00    | $34.43     | $0.55     | 342                                             | 3.598e+07
5    | Marc Beniof
```

**Question 4.2.** How many CEOs made more than $30 million in total pay? Find the value using code, then check that the value you found is consistent with what you see in the histogram.


```python
num_ceos_more_than_30_million_2 = compensation.select("Total Pay ($)", are.above(30))
num_ceos_more_than_30_million_2
```

Output:
```
File /usr/local/python/3.12.1/lib/python3.12/site-packages/datascience/tables.py:2171, in Table._as_labels(self, label_or_labels)
   2169 def _as_labels(self, label_or_labels):
   2170     """Convert single label to list and convert indices to labels."""
-> 2171     return [self._as_label(s) for s in _as_labels(label_or_labels)]

File /usr/local/python/3.12.1/lib/python3.12/site-packages/datascience/tables.py:2167, in Table._as_label(self, index_or_label)
   2165     return self.labels[index_or_label]
   2166 else:
-> 2167     raise ValueError(str(index_or_label) + ' is not a label or index')

ValueError: <datascience.predicates._combinable object at 0x7e8f02a5e900> is not a label or index
```

```python
grader.check("q42")
```

Output:
```
q42 results:
    q42 - 1 result:
        ❌ Test case failed
        Trying:
            num_ceos_more_than_30_million_2 == 5
        Expecting:
            True
        **********************************************************************
        Line 1, in q42 0
        Failed example:
            num_ceos_more_than_30_million_2 == 5
        Expected:
            True
        Got:
            False
```

```python
# Run this cell to run all the autograde tests, then use Github to export to HTML.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...


Your submission received the following results when run against available test cases:

    q11 results: All test cases passed!

    q12 results: All test cases passed!

    q13 results: All test cases passed!

    q14 results:
        q14 - 1 result:
            ❌ Test case failed
            Trying:
                print_kth_top_movie_year(4)
            Expecting:
                Year number 4 for total gross movie sales was: 2009
            **********************************************************************
            Line 1, in q14 0
            Failed example:
                print_kth_top_movie_year(4)
            Expected:
                Year number 4 for total gross movie sales was: 2009
            Got:
                Year number 4 for total gross movie sales was 2014

    q21 results: All test cases passed!

    q22 results: All test cases passed!

    q23 results: All test cases passed!

    q24 results: All test c
```

Tests for q11 (points: None)
```
>>> # Make sure your function has the proper syntax!
>>> to_percentage(.35) == 35.0
True
```
```
>>> # Make sure you are passing in the proper argument!
>>> twenty_percent == 20.0
True
```


Tests for q12 (points: None)
```
>>> # Make sure you're using a_proportion!
>>> np.isclose(a_percentage, 70.71067811865476)
True
```


Tests for q13 (points: None)
```
>>> num_non_vowels("Go bears!") == 6
True
```


Tests for q14 (points: None)
```
>>> print_kth_top_movie_year(4)
Year number 4 for total gross movie sales was: 2009
```


Tests for q21 (points: None)
```
>>> # Make sure you are examining the values in the column, not the column itself
>>> import numpy
>>> total_pay_type != numpy.ndarray
True
```
```
>>> 'str' in str(total_pay_type)
True
```
```
>>> isinstance(total_pay_type, type)
True
```
```
>>> # Make sure to call the type function on a value in the column
>>> total_pay_type != int
True
```


Tests for q22 (points: None)
```
>>> str(mark_hurd_pay_string)
'$53.25 '
```


Tests for q23 (points: None)
```
>>> # Your answer should be a number
>>> type(mark_hurd_pay) != str
True
```
```
>>> # Don't forget to give your answer in dollars, not millions of 
>>> # Dollars! 
>>> mark_hurd_pay != 5325
True
```
```
>>> # Don't forget to give your answer in dollars, not millions of 
>>> # Dollars! 
>>> mark_hurd_pay == 53250000
True
```


Tests for q24 (points: None)
```
>>> convert_pay_string_to_number("$100 ") == 100000000.0
True
```
```
>>> convert_pay_string_to_number("$23 ") == 23000000.0
True
```


Tests for q31 (points: None)
```
>>> len(some_functions)
3
```
```
>>> # The first thing in your array may not be a function
>>> callable(some_functions.item(0))
True
```
```
>>> # The second thing in your array may not be a function
>>> callable(some_functions.item(1))
True
```
```
>>> # The third thing in your array may not be a function
>>> callable(some_functions.item(2))
True
```


Tests for q32 (points: None)
```
>>> # You either didn't add the 'Total Pay ($)' column, 
>>> # or you mislabeled it
>>> 'Total Pay ($)' in compensation.labels
True
```
```
>>> # You have the column in your table, but the values may be wrong
>>> t = compensation.sort('Total Pay ($)', descending = True)
>>> np.isclose(t.column('Total Pay ($)').item(0), 53250000.0)
True
```


Tests for q33 (points: None)
```
>>> import math
>>> math.isclose(average_total_pay, 11558613.861386139, rel_tol = 0.1)
True
```


Tests for q34 (points: None)
```
>>> len(cash_proportion) == 101
True
```
```
>>> import math
>>> math.isclose(cash_proportion.item(0), 0.01784038, rel_tol = .001)
True
```


Tests for q35 (points: None)
```
>>> # Make sure to remove the "(No previous year)" CEOs 
>>> "(No previous year)" not in with_previous_compensation.column("% Change")
True
```
```
>>> import math
>>> # You have the column, but some of your values may be wrong
>>> t = with_previous_compensation.sort("2014 Total Pay ($)", descending = True)
>>> value = t.column("2014 Total Pay ($)").item(0)
>>> math.isclose(value, 67700000.0, rel_tol = 0.01)
True
```
```
>>> # You have the column, but your number of rows is off
>>> with_previous_compensation.num_rows == 80
True
```


Tests for q36 (points: None)
```
>>> import math
>>> math.isclose(average_pay_2014, 11794790.817048479, rel_tol = 0.01)
True
```


Tests for q42 (points: None)
```
>>> num_ceos_more_than_30_million_2 == 5
True
```
